---
title: 'Population vs. Sample'
description: 'When it''s too expensive to look at all the data'
---

## How good is sampling?

```yaml
type: NormalExercise
key: cc3c7b4450
xp: 100
```

Sampling is a way to save resources by looking at part of the population instead of all of it. How good is it?

`@instructions`
You will be given a list of numerical values in an array called `values`. Your job is to create graphs to illustrate how well sampling works with different sample sizes:

1. Create all the possible samples for each sample size
1. For each sample, calculate the central tendencies (mean and median) and standard deviation
1. Create graphs of sample sizes and the difference between the calculated values and the values for the entire population

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
```
